# Strings

Topics:
1. How to get a string
2. String interpolation
3. String concatenation

## How to get a string

Enclose your characters in " " or """ """!

```julia
s1 = "I am a string."
```

Output:
```
"I am a string."
```

```julia
typeof(s1[1])
```

Output:
```
Char
```

```julia
typeof(s1[end])
```

Output:
```
Char
```

```julia
s1 = "I'm string."
```

Output:
```
"I'm string."
```

```julia
s2 = """I am also a string. """
typeof(s2)
```

Output:
```
String
```

There are a couple functional differences between strings enclosed in single and triple quotes. <br>
One difference is that, in the latter case, you can use quotation marks within your string.

```julia
println("""Here, we get an "error" because it's ambiguous where this string ends """)
```

Output:
```
Here, we get an "error" because it's ambiguous where this string ends
```

```julia

```

```julia
println("""Here, we get an error because it's ambiguous where this string ends""")
```

Output:
```
Here, we get an error because it's ambiguous where this string ends
```

```julia

```

```julia

```

```julia
println("""Look, Mom, no "errors!!! """)
```

Output:
```
Look, Mom, no "errors!!!
```

Note that ' ' define a character, but NOT a string!

```julia
typeof('a')
```

```julia
'We will get an error here'
```

## String interpolation

We can use the $ sign to insert existing variables into a string and to evaluate expressions within a string. <br>
Below is an example that contains some highly sensitive personal information.

```julia
s2 + s1 # this is for python
```

Output:
```
Stacktrace:
 [1] top-level scope
   @ In[16]:1
```

```julia
s1*s2 # correct way to do concatention for Julia
```

Output:
```
"I'm string.I am also a string. "
```

```julia
string(s1,"  ",s2)
```

Output:
```
"I'm string.  I am also a string. "
```

```julia
name = "Jane"
num_fingers = 10
num_toes = 10
```

Output:
```
10
```

```julia
println("Hello, my name is " * name, " hello")

```

Output:
```
Hello, my name is Jane hello
```

```julia
println("Hello, my name is $name.")
println("I have $num_fingers fingers and $num_toes toes.")
```

Output:
```
Hello, my name is Jane.
I have 10 fingers and 10 toes.
```

```julia
println("That is $(num_fingers + num_toes) digits in all!!")
```

Output:
```
That is 20 digits in all!!
```

```julia
x="That is $(num_fingers + num_toes) digits in all!!"
x
```

Output:
```
"That is 20 digits in all!!"
```

## String concatenation

Below are three ways we can concatenate strings! <br><br>
The first way is to use the `string()` function. <br>
`string()` converts non-string inputs to strings.

```julia
s3 = "How many cats ";
s4 = "is too many cats?";
😺 = 10
```

Output:
```
10
```

```julia
string(s3, s4)
```

Output:
```
"How many cats is too many cats?"
```

```julia
string("I don't know, but ", 😺,"  is too few.")
```

Output:
```
"I don't know, but 10  is too few."
```

```julia
println("I don't know, but , $😺,  is too few.")
```

Output:
```
I don't know, but , 10,  is too few.
```

We can also use `*` for concatenation!

```julia
s3*s4
```

Output:
```
"How many cats is too many cats?"
```

```julia

```

## Input from Users
We can capture or read user inputs from console in Julia. This can be done through inbuilt I/O methods like :

- readline()
- readlines()

```julia
school_name = readline()
println(school_name)
```

Output:
```
stdin>  UH
UH
```

```julia
print("What's the name of your school?\n\n")
flush(stdout)
school_name = readline()
println(school_name)
```

Output:
```
Whas the name of your school?

stdin>  UH
UH
```

```julia
println("What's the name of your school?")
flush(stdout)
school_name = readline()
println("I attend $school_name .")
```

Output:
```
What's the name of your school?
stdin>  uh
I attend uh .
```

```julia
# Julia program to take input from user 

# prompt to input 
print("What's your name ? \n\n") 
flush(stdout)
# Calling rdeadline() function 
name = readline() 
flush(stdout)

println("The name is ", name) 
print("\n\n") 

# typeof() determines the datatype. 
println("Type of the input is: ", typeof(name)) 

```

Output:
```
What's your name ? 

stdin>  jerry
The name is jerry


Type of the input is: String
```

### Reading numerical data types from console

Julia can also read inputs as numbers and make use of them in further computations. This is done using the parse() method, using which we can convert a numeric string(Float or Int) into a numerical value.

```julia
# Julia program to calculate Area of Circle
# 1 float obtained from console/user input 


# Prompt to enter 
println("""Enter radius (floats preferred)""") 
flush(stdout)

# Taking Input from user 

r = readline() 
r = parse(Float64, r) 
area = pi * r^2
println("The area of circle with radius $r is ", area) 
 


```

Output:
```
Enter radius (floats preferred)
stdin>  67.0
The area of circle with radius 67.0 is 14102.609421964582
```

### Reading N lines of input from stdin or file
- N-lines of input can be captured and stored within a String Array.

#### Reading from files

```julia
#Lets get a file from the web
```

```julia

;curl https://www.gutenberg.org/files/74/74-0.txt -o TomSawyer.txt
```

Output:
```
% Total    % Received % Xferd  Average Speed   Time    Time     Time  Current
                                 Dload  Upload   Total   Spent    Left  Speed
100  396k  100  396k    0     0   697k      0 --:--:-- --:--:-- --:--:--  698k
```

```julia
;head TomSawyer.txt
```

Output:
```
﻿*** START OF THE PROJECT GUTENBERG EBOOK THE ADVENTURES OF TOM SAWYER ***




THE ADVENTURES OF TOM SAWYER


By Mark Twain
```

```julia
# Julia program to take 
# multi-lined input from user
println("Enter multi-lined text, press Ctrl-D when done") 
```

Output:
```
Enter multi-lined text, press Ctrl-D when done
```

```julia
# Calling readlines() function 
lines=readlines("TomSawyer.txt");
```

```julia

# Calling readlines() function 
lines = readlines("TomSawyer.txt") 

println("total no.of.lines read: ", length(lines)) 

#println(lines) 

# Getting type of Input values 
println("type of input: ", typeof(lines)) 

```

Output:
```
total no.of.lines read: 8894
type of input: Vector{String}
```

```julia
#read the file in one gulp and store in string

file = open("TomSawyer.txt")
contents= read(file,String)
close(file)
```

```julia
#print the contents
print(contents)
```

```julia
#also, you can just read file directly with read command only
contents2=read("TomSawyer.txt", String);
```

```julia
#print the contents
print(contents2)
```

#### Writing to files

```julia
# first lets write to a file real quick
file_io=open("myfile.txt","w")
mystring="""It a whole new world out there!!!
Rise and Shine"""
write(file_io, mystring)
close(file_io)
```

```julia
#for Mac/Unix/Linux users
```

```julia
;cat myfile.txt
```

Output:
```
It a whole new world out there!!!
Rise and Shine
```

```julia
#for Mac/Unix/Linux users
```

```julia
#run(`cmd /c type myfile.txt`)
```

```julia

```

```julia

```

```julia
contents=read("myfile.txt", String)
println(String(contents))
```

Output:
```
It a whole new world out there!!!
Rise and Shine
```

```julia

```

```julia
lines2=readlines("myfile.txt")
println(lines2)
```

Output:
```
["It a whole new world out there!!!", "Rise and Shine"]
```

```julia
println(lines2[2])
```

```julia
println(lines2[1])
```

```julia
typeof(lines2)
```

```julia

```

Output:
```
search: LazySt[0m[1mr[22ming [0m[1mr[22m[0m[1me[22m[0m[1ma[22m[0m[1md[22mlink SubSt[0m[1mr[22ming [0m[1mr[22m[0m[1me[22m[0m[1ma[22m[0m[1md[22mline St[0m[1mr[22ming st[0m[1mr[22ming [0m[1mr[22m[0m[1me[22m[0m[1ma[22m[0m[1md[22mline[0m[1ms[22m

Couldn't find [36mreadString[39m
Perhaps you meant String
  No documentation found.

  Binding [36mreadString[39m does not exist.
```

### Exercises

#### 2.1 
Create a string that says "hi" 1000 times, first with `repeat` and then with the exponentiation operator, which can call `*` under the hood. Assign it the variable `hi` below.

<!-- grade_id: cell-5b38bdb1115161b4, points: 0, grade: false, solution: true -->
```julia

```

```julia

```

```julia

```

```julia

```

<!-- grade_id: cell-20b1027a968fbb48, points: 1, grade: true, solution: false -->
```julia
@assert hi == "hihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihi"
```

#### 2.2 
Declare two variables

```julia
a = 3
b = 4
```
and use them to create two strings:
```julia
"3 + 4"
"7" 
```
and store the results in `c` and `d` respectively

<!-- grade_id: cell-36beb524c8ecd33b, points: 0, grade: false, solution: true -->
```julia
 
```

```julia

```

```julia
 
```

<!-- grade_id: cell-14c3c5b015a97892, points: 1, grade: true, solution: false -->
```julia
@assert c == "3+4"
@assert d == "7"
```

Please run the @assert cells once you are done with the exercises to validate your answer.