# Basic linear algebra in Julia
Originally started by : Andreas Noack Jensen
with edits from Jane Herriman

```julia
using Random
Random.seed!(1234)
```

First let's define a random matrix

```julia
A = rand(1:9,3,3)
```

Define a vector of ones

```julia
x = fill(1.0, (3,)) # = fill(1.0, 3)
```

Notice that $A$ has type Array{Int64,2} but $x$ has type Array{Float64,1}. Julia defines the aliases Vector{Type}=Array{Type,1} and Matrix{Type}=Array{Type,2}. 

Many of the basic operations are the same as in other languages
 

```julia

```

## Generate random numbers in a range
- use case<br>
*Develop a simple game that teaches kindergartners how to add single-digit numbers. Your function game() will take an integer n as input and then ask n single-digit addtion questions. The numbers to be added should be chosen randomly from the range [1, 20] (i.e.,1 to 20 inclusive). The user will enter the answer when prompted. Your function should print 'Correct' for correct answers and 'Incorrect' for incorrect answers. After n questions, your function should print the number of correct answers.*

```julia
function add_game(x::Int)
    count=0
    for i in 1:x
      a=rand(1:20,1)[1];
      b=rand(1:20,1)[1];
      println("What is  $a + $b")
      flush(stdout)
      ans=parse(Int64, readline())
      if ans == (a + b)
            println("correct answer")
            count+=1 
       else
            println("wrong answer")
       end
    end
        println("You scored $count out of $x points")
end

```

```julia
add_game(3)
```

 #### Multiplication

```julia
A 
 
```

```julia
 
x
 
```

```julia
b = A*x
```

#### Transposition
As in other languages `A'` is the conjugate transpose, or adjoint

```julia
A
```

```julia
A'
```

and we can get the transpose with

```julia
transpose(A)
```

#### Transposed multiplication
Julia allows us to write this without *

```julia
A'A
```

#### Solving linear systems 
The problem $Ax=b$ for ***square*** $A$ is solved by the \ function.

```julia
A\b
```

`A\b` gives us the *least squares solution* if we have an overdetermined linear system (a "tall" matrix)

```julia
Atall = rand(3, 2)
```

```julia
Atall\b
```

and the *minimum norm least squares solution* if we have a rank-deficient least squares problem

```julia
v = rand(3)
```

### Horizontal concatention

```julia

rankdef = hcat(v, v)
```

```julia
rankdef\b
```

Julia also gives us the minimum norm solution when we have an underdetermined solution (a "short" matrix)

```julia
bshort = rand(2)
Ashort = rand(2, 3)
```

```julia
Ashort\bshort
```

- Solving simultaneous equations

4x + 3y = 10

x + y = 3

```julia
# write coefficient data in matrix form
A = [[4 3];[1 1]]
```

```julia
#Write results in vector form
b = [10,3]
```

```julia
#then solve
x = A\b
```

# The LinearAlgebra library

While much of linear algebra is available in Julia by default (as shown above), there's a standard library named `LinearAlgebra` that brings in many more relevant names and functions. In particular, it provides factorizations and some structured matrix types.  As with all packages, you can bring these additional features into your session with a `using LinearAlgebra`.

```julia
using Pkg
```

```julia

Pkg.add("LinearAlgebra")
using LinearAlgebra
```

To list the various functions available in the package we use the `names` command

```julia
for i in names(LinearAlgebra)
    print(i, " ")
end
```

```julia

```

To gather information of all the current variables in the environment use `varinfo`

```julia
varinfo()
```

#### How do I delete an object in memory? ####

Julia does not have an analog of MATLAB's `clear` function; once a name is defined in a Julia session (technically, in module `Main`), it is always present.

If memory usage is your concern, you can always replace objects with ones that consume less memory. For example, if `A` is a gigabyte-sized array that you no longer need, you can free the memory with `A = nothing`. The memory will be released the next time the garbage collector runs; you can force this to happen with `gc()`. Moreover, an attempt to use `A` will likely result in an error, because most methods are not defined on type `Nothing`.

#### How can I modify the declaration of a type in my session?

Perhaps you've defined a type and then realize you need to add a new field. If you try this, you get the error:

`ERROR: invalid redefinition of constant MyType`

Types in module `Main` cannot be redefined.

```julia
A=nothing
Ashort=nothing
Atall=nothing
rankdef=nothing
```

```julia
GC.gc()
```

```julia
varinfo()
```

### Exercises

```julia
v=[0.4913831957970459
 0.5651453592612876
 0.2538117862083361]
```

```julia

```

 

#### 11.1a 
Take the inner product (or "dot" product) of a vector `v` with itself and assign it to variable `v_squared`.



```julia

```

#### 11.1b
Take the inner product (or "dot" product) of a vector `v` with itself and assign it to variable `dot_v`. Use the `dot` function



```julia
using LinearAlgebra

```

```julia
 
```

```julia
 
```

```julia

```

```julia
@assert dot_v ==  0.6252671450245539 
```

#### 11.2 
Take the outer product of a vector v with itself and assign it to variable `cross_v`

```julia

```

```julia

```

<!-- grade_id: cell-e6b6970ffe104df5, points: 1, grade: true, solution: false -->
```julia
@assert cross_v == [0, 0, 0]
```

```julia

```

Please execute the @assert cell once you are done with the exercises to validate your solution.